# Q2

This question will walk you through the process of how a coding assignment in Jupyter will work.



There's one question below--a small code snippet. Don't worry about what it does; instead, focus on the **double-quotes**.

See how they're right next to each other? Put something between those quotes. Anything at all. The autograder for this question will literally look to see if there's so much as a single blank space between the quotes, and will call the answer correct. `" "` would be correct; so would `"this is a lot more than just a space"`. Make something up! But whatever you do, *don't delete the quotes*.

 1. Copy the *whole* line that looks like: `instring = ""`
 2. Paste it between the `BEGIN SOLUTION` and `END SOLUTION`
 3. Put something inside the double-quotes! Anything!
 
**Step 2 is crucial, otherwise your code will fail!**

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
instring = ""

### BEGIN SOLUTION
instring = "Today is a nice day"
### END SOLUTION

# IGNORE EVERYTHING BELOW THIS LINE!!! THANKS!!!

def things_and_stuff(some_string):
    return some_string

out = things_and_stuff(instring)
print("Here's what you wrote: {}".format(out))
```

Output:
```
Here's what you wrote: Today is a nice day
```

<!-- grade_id: q4a_tests, points: 20, grade: true, solution: false -->
```python
assert type(out) == str
assert len(out) > 0
```

```python

```